<a href="https://colab.research.google.com/github/jddantes/Seinfeld/blob/master/01_numpy_exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joseph Daniel Dantes"
COLLABORATORS = ""
```

---

## Assorted Numpy Exercises

#### 1. Import the numpy package under the name `np` (★☆☆)

`hint: import … as`

<!-- grade_id: cell-0cff6875083d6b25, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
# raise NotImplementedError()
```

<!-- grade_id: cell-6b0de2ec8bbb7be8, points: 1, grade: true, solution: false -->
```python
try:
  np.array([0., 0., 0.])
  assert True
except NameError:
  assert False
```

#### 2. Print the numpy version and the configuration (★☆☆)

`hint: np.__version__, np.show_config)`

<!-- grade_id: cell-bda094f852c623ab, points: 0, grade: false, solution: true -->
```python
version = None
# YOUR CODE HERE
version = np.__version__
# raise NotImplementedError()
```

<!-- grade_id: cell-1a1abe159355bd8f, points: 1, grade: true, solution: false -->
```python
# Colab-Dependent 8)
assert version == '1.25.2'
```

#### 3. Create a null vector of size 10 (★☆☆)

`hint: np.zeros`

<!-- grade_id: zeros, points: 0, grade: false, solution: true -->
```python
answer3 = None
# YOUR CODE HERE
answer3 = np.zeros(10)
# raise NotImplementedError()
```

<!-- grade_id: zeros-test, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(answer3, np.array([0., 0., 0., 0., 0., 0., 0., 0., 0., 0.]))
```

#### 4. How to find the memory size of any array (★☆☆)

`hint: size, itemsize`

<!-- grade_id: cell-27bb53e82226d8d6, points: 0, grade: false, solution: true -->
```python
Z = np.zeros((10,10))
itemsize = None
itemsize = Z.itemsize * Z.size
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-6b241446992eb73a, points: 1, grade: true, solution: false -->
```python
assert itemsize == 800
```

#### 5. How to get the documentation of the numpy add function from the command line? (★☆☆)

`hint: np.info`

<!-- grade_id: cell-be69e06c4c948795, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
np.info(np.add)
# raise NotImplementedError()
```

Output:
```
add(x1, x2, /, out=None, *, where=True, casting='same_kind', order='K', dtype=None, subok=True[, signature, extobj])

Add arguments element-wise.

Parameters
----------
x1, x2 : array_like
    The arrays to be added.
    If ``x1.shape != x2.shape``, they must be broadcastable to a common
    shape (which becomes the shape of the output).
out : ndarray, None, or tuple of ndarray and None, optional
    A location into which the result is stored. If provided, it must have
    a shape that the inputs broadcast to. If not provided or None,
    a freshly-allocated array is returned. A tuple (possible only as a
    keyword argument) must have length equal to the number of outputs.
where : array_like, optional
    This condition is broadcast over the input. At locations where the
    condition is True, the `out` array will be set to the ufunc result.
    Elsewhere, the `out` array will retain its original value.
    Note that if an uninitialized `out` array is created via the default
    ``out
```

#### 6. Create a null vector of size 10 but the fifth value which is 1 (★☆☆)

`hint: array[4]`

<!-- grade_id: cell-84c9807e3eab5bdb, points: 0, grade: false, solution: true -->
```python
Z_ans = None
# YOUR CODE HERE
Z_ans = np.zeros(10)
Z_ans[4] = 1
# raise NotImplementedError()
```

<!-- grade_id: cell-280ddaccfb592b37, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(Z_ans, np.array([0 if i!=4 else 1 for i in range(10)]))
```

#### 7. Create a vector with values ranging from 10 to 49 (★☆☆)

`hint: arange`

<!-- grade_id: cell-fae07e150c7204ae, points: 0, grade: false, solution: true -->
```python
answer7 = None
# YOUR CODE HERE
answer7 = np.array(range(10, 50))
# raise NotImplementedError()
```

<!-- grade_id: cell-d35392b97b5ff0ca, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(answer7, np.array([i for i in range(10, 50)]))
```

#### 9. Create a 3x3 matrix with values ranging from 0 to 8 (★☆☆)

`hint: reshape`

<!-- grade_id: cell-972343f87432e5b3, points: 0, grade: false, solution: true -->
```python
answer9 = np.arange(0, 9)
# YOUR CODE HERE
answer9 = answer9.reshape((3, 3))
# raise NotImplementedError()
```

<!-- grade_id: cell-53a6f8c1396334ca, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(answer9, np.matrix([[x for x in range(i*3, i*3+3)] for i in range(3)]))
```

#### 10. Find indices of non-zero elements from [1,2,0,0,4,0] (★☆☆)

`hint: np.nonzero`

<!-- grade_id: cell-87813f9914f2e2fe, points: 0, grade: false, solution: true -->
```python
nz = None
# YOUR CODE HERE
tmp = np.array([1, 2, 0, 0, 4, 0])

nz = tmp.nonzero()
# raise NotImplementedError()
```

<!-- grade_id: cell-8a41054e4ea0161f, points: 1, grade: true, solution: false -->
```python
assert np.array_equal([0, 1, 4], nz[0])
```

#### 11. Create a 3x3 identity matrix (★☆☆)

`hint: np.eye`

<!-- grade_id: cell-55dabd92fbdaa065, points: 0, grade: false, solution: true -->
```python
Z = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b88758a16d55c65a, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(Z, np.array([1 if i%4 == 0 else 0 for i in range(9)]).reshape((3,3)))
```

#### 12. Create a 3x3x3 array with random values (★☆☆)

`hint: np.random.random`

<!-- grade_id: cell-44e4bee1a203d74a, points: 0, grade: false, solution: true -->
```python
np.random.seed(10)
Z = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-14b700666966532d, points: 1, grade: true, solution: false -->
```python
exp = np.array([[[0.77132064, 0.02075195, 0.63364823],
        [0.74880388, 0.49850701, 0.22479665],
        [0.19806286, 0.76053071, 0.16911084]],
       [[0.08833981, 0.68535982, 0.95339335],
        [0.00394827, 0.51219226, 0.81262096],
        [0.61252607, 0.72175532, 0.29187607]],
       [[0.91777412, 0.71457578, 0.54254437],
        [0.14217005, 0.37334076, 0.67413362],
        [0.44183317, 0.43401399, 0.61776698]]])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
```

#### 13. Create a 10x10 array with random values and find the minimum and maximum values (★☆☆)

`hint: min, max`

<!-- grade_id: cell-7e58389b5fb99379, points: 0, grade: false, solution: true -->
```python
np.random.seed(10)
Z = None
Zmin, Zmax = None, None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c910f6b519aedb2f, points: 1, grade: true, solution: false -->
```python
assert round(Zmin, 8) == 0.00394827
assert round(Zmax, 8) == 0.98762547
```

#### 14. Create a random vector of size 30 and find the mean value (★☆☆)

`hint: mean`

<!-- grade_id: cell-62066049ca50fc97, points: 0, grade: false, solution: true -->
```python
np.random.seed(10)
Z = None
Zmean = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fccccdc671b10a03, points: 1, grade: true, solution: false -->
```python
assert round(Zmean, 8) == 0.51100906
```

#### 15. Create a 10x10 2d array with 1 on the border and 0 inside (★☆☆)

> Indented block



`hint: array[1:-1, 1:-1]`

<!-- grade_id: cell-39c98695b80dfd88, points: 0, grade: false, solution: true -->
```python
array = None
# YOUR CODE HERE
raise NotImplementedError()
```

```python
array
```

<!-- grade_id: cell-d277d27f07e61c2c, points: 1, grade: true, solution: false -->
```python
exp = np.array([[1., 1., 1., 1., 1., 1., 1., 1., 1., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 0., 0., 0., 0., 0., 0., 0., 0., 1.],
       [1., 1., 1., 1., 1., 1., 1., 1., 1., 1.]])
assert np.array_equal(array, exp)
```

#### 16. How to add a border (filled with 0's) around an existing array? (★☆☆)

`hint: np.pad`

<!-- grade_id: cell-a4083cc764aa0701, points: 0, grade: false, solution: true -->
```python
Z = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2ded8e2a73e101cc, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([[0., 0., 0., 0., 0., 0., 0.],
               [0., 1., 1., 1., 1., 1., 0.],
               [0., 1., 1., 1., 1., 1., 0.],
               [0., 1., 1., 1., 1., 1., 0.],
               [0., 1., 1., 1., 1., 1., 0.],
               [0., 1., 1., 1., 1., 1., 0.],
               [0., 0., 0., 0., 0., 0., 0.]])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 17. What is the result of the following expression? (★☆☆)
```python
0 * np.nan
np.nan == np.nan
np.inf > np.nan
np.nan - np.nan
np.nan in set([np.nan])
0.3 == 3 * 0.1
```

`hint: NaN = not a number, inf = infinity`

<!-- grade_id: cell-72ab815a75c8c211, points: 0, grade: false, solution: true -->
```python
A = None
B = None
C = None
D = None
E = None
F = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec04eb0ff109bda1, points: 1, grade: true, solution: false -->
```python

```

#### 18. Create a 5x5 matrix with values 1,2,3,4 just below the diagonal (★☆☆)

`hint: np.diag`

<!-- grade_id: cell-9a0d1cb98d50ad4e, points: 0, grade: false, solution: true -->
```python
Z = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-07b6d4d3a4c333ec, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([[0, 0, 0, 0, 0],
               [1, 0, 0, 0, 0],
               [0, 2, 0, 0, 0],
               [0, 0, 3, 0, 0],
               [0, 0, 0, 4, 0]])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 21. Create a checkerboard 8x8 matrix using the tile function (★☆☆)

`hint: np.tile`

<!-- grade_id: cell-ff6d52fa8da1b789, points: 0, grade: false, solution: true -->
```python
Z = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-72f686ddb377d7d9, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([[0, 1, 0, 1, 0, 1, 0, 1],
               [1, 0, 1, 0, 1, 0, 1, 0],
               [0, 1, 0, 1, 0, 1, 0, 1],
               [1, 0, 1, 0, 1, 0, 1, 0],
               [0, 1, 0, 1, 0, 1, 0, 1],
               [1, 0, 1, 0, 1, 0, 1, 0],
               [0, 1, 0, 1, 0, 1, 0, 1],
               [1, 0, 1, 0, 1, 0, 1, 0]])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 24. Multiply a 5x3 matrix by a 3x2 matrix (real matrix product) when both initially contain 1 (★☆☆)

`No hints provided...`

<!-- grade_id: cell-d6ab040c4bffe8a4, points: 0, grade: false, solution: true -->
```python
Z = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-89d59c9ae96b8129, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([[3., 3.],
               [3., 3.],
               [3., 3.],
               [3., 3.],
               [3., 3.]])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 25. Given a 1D array, negate all elements which are between 3 and 8, in place. (★☆☆)

`hint: > or <`

<!-- grade_id: cell-2c3164e3524c0b09, points: 0, grade: false, solution: true -->
```python
Z = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-32f0a60de00bb0db, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([ 0,  1,  2,  3, -4, -5, -6, -7,  8,  9, 10])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 26. What is the output of the following script? (★☆☆)
```python
# Author: Jake VanderPlas

print(sum(range(5),-1))
from numpy import *
print(sum(range(5),-1))
```

`hint: np.sum`

<!-- grade_id: cell-0467f751410b0893, points: 0, grade: false, solution: true -->
```python
print1 = None
print2 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-568b0398ea57c5eb, points: 1, grade: true, solution: false -->
```python

```

#### 27. Consider an integer vector Z, which of these expressions are legal? (★☆☆)
```python
Z**Z
2 << Z >> 2
Z <- Z
1j*Z
Z/1/1
Z<Z>Z
```

`No hints provided...`

<!-- grade_id: cell-4f6b39b8503069b7, points: 0, grade: false, solution: true -->
```python
# Put `True` if legal otherwise, put `False`.

A = None # Z**Z
B = None # 2 << Z >> 2
C = None # Z <- Z
E = None # 1j*Z
E = None # Z/1/1
F = None # Z<Z>Z

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-67fded3e84618e95, points: 1, grade: true, solution: false -->
```python

```

#### 28. What are the result of the following expressions? (★☆☆)
```python
np.array(0) / np.array(0)
np.array(0) // np.array(0)
np.array([np.nan]).astype(int).astype(float)
```

`No hints provided...`

<!-- grade_id: cell-b92fd082c5465ced, points: 0, grade: false, solution: true -->
```python
A = None
B = None
C = None

# YOUR CODE HERE
raise NotImplementedError()
```

```python
C
```

<!-- grade_id: cell-2a958062841bd93c, points: 1, grade: true, solution: false -->
```python

```

#### 30. How to find common values between two arrays? (★☆☆)

`hint: np.intersect1d`

<!-- grade_id: cell-e56f557c9cb773ba, points: 0, grade: false, solution: true -->
```python
np.random.seed(10)
Z1 = np.random.randint(0,10,10)
Z2 = np.random.randint(0,10,10)

ans_30 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ae809cda496c7561, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([0, 1, 4, 8])

assert np.array_equal(np.matrix.round(ans_30, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 31. How to ignore all numpy warnings (not recommended)? (★☆☆)

`hint: np.seterr, np.errstate`

<!-- grade_id: cell-821f3dc3780f8562, points: 0, grade: false, solution: true -->
```python

Z = np.ones(1) / 0
# YOUR CODE HERE
raise NotImplementedError()
```

#### 32. Is the following expressions true? (★☆☆)
```python
np.sqrt(-1) == np.emath.sqrt(-1)
```

`hint: imaginary number`

<!-- grade_id: cell-88fe5d7371e493a5, points: 0, grade: false, solution: true -->
```python
Z = None # True or False
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-77948163b3b1f949, points: 1, grade: true, solution: false -->
```python

```

#### 35. How to compute ((A+B)*(-A/2)) in place (without copy)? (★★☆)

`hint: np.add(out=), np.negative(out=), np.multiply(out=), np.divide(out=)`

<!-- grade_id: cell-6078ef2ebe6681b7, points: 0, grade: false, solution: true -->
```python
A = np.ones(3)*1
B = np.ones(3)*2

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e6aa5917ec4a15c, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([-1.5, -1.5, -1.5])

assert np.array_equal(np.matrix.round(A, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 36. Extract the integer part of a random array of positive numbers using 4 different methods (★★☆)

`hint: %, np.floor, astype, np.trunc`

<!-- grade_id: cell-3e073effc35d1666, points: 0, grade: false, solution: true -->
```python
np.random.seed(10)
ans_36 = np.random.uniform(0,10,10)

modulo_ans = None
floor_ans = None
type_ans = None
truc_ans = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0d80b94ebfa9f351, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
modulo_exp = np.array([0.85359818, 0.53393346, 0.03948266, 0.12192263, 0.12620962,
       0.12526067, 0.21755317, 0.91876068, 0.17774123, 0.14575783])
floor_exp = np.array([6., 9., 0., 5., 8., 6., 7., 2., 9., 7.])
type_exp = np.array([6., 9., 0., 5., 8., 6., 7., 2., 9., 7.])
truc_exp = np.array([6, 9, 0, 5, 8, 6, 7, 2, 9, 7])

assert np.array_equal(np.matrix.round(modulo_ans, 6), np.matrix.round(modulo_exp, 6))
assert np.array_equal(np.matrix.round(floor_ans, 6), np.matrix.round(floor_exp, 6))
assert np.array_equal(np.matrix.round(type_ans, 6), np.matrix.round(type_exp, 6))
assert np.array_equal(np.matrix.round(truc_ans, 6), np.matrix.round(truc_exp, 6))
### END TEST
```

#### 37. Create a 5x5 matrix with row values ranging from 0 to 4 (★★☆)

`hint: np.arange`

<!-- grade_id: cell-e7b56df22efdadcb, points: 0, grade: false, solution: true -->
```python
Z = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d25608ec91eab549, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([[0., 1., 2., 3., 4.],
       [0., 1., 2., 3., 4.],
       [0., 1., 2., 3., 4.],
       [0., 1., 2., 3., 4.],
       [0., 1., 2., 3., 4.]])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```

#### 40. Create a random vector of size 10 and sort it (★★☆)

`hint: sort`

<!-- grade_id: cell-87e13b86256435fe, points: 0, grade: false, solution: true -->
```python
np.random.seed(10)
Z = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5ecf24d586341ea0, points: 1, grade: true, solution: false -->
```python
### BEGIN TEST
exp = np.array([0.02075195, 0.08833981, 0.16911084, 0.19806286, 0.22479665,
       0.49850701, 0.63364823, 0.74880388, 0.76053071, 0.77132064])

assert np.array_equal(np.matrix.round(Z, 6), np.matrix.round(exp, 6))
### END TEST
```